<h1>Table of Contents<span class="tocSkip"></span></h1>
<div class="toc"><ul class="toc-item"><li><span><a href="#Lab-4:-Unicode,-Multiple-Ways,-and-Lambda-Functions" data-toc-modified-id="Lab-4:-Unicode,-Multiple-Ways,-and-Lambda-Functions-1">Lab 4: Unicode, Multiple Ways, and Lambda Functions</a></span></li><li><span><a href="#Learning-Outcomes" data-toc-modified-id="Learning-Outcomes-2">Learning Outcomes</a></span><ul class="toc-item"><li><ul class="toc-item"><li><span><a href="#By-the-end-of-this-notebook,-you-should-be-able-to:" data-toc-modified-id="By-the-end-of-this-notebook,-you-should-be-able-to:-2.0.1">By the end of this notebook, you should be able to:</a></span></li></ul></li></ul></li><li><span><a href="#Always-a-string" data-toc-modified-id="Always-a-string-3">Always a string</a></span></li><li><span><a href="#Messed-up-content" data-toc-modified-id="Messed-up-content-4">Messed up content</a></span></li><li><span><a href="#ELT-(Extract,-Transform,-Load):-A-text-file-of-emojis-into-a-Python-dict" data-toc-modified-id="ELT-(Extract,-Transform,-Load):-A-text-file-of-emojis-into-a-Python-dict-5">ELT (Extract, Transform, Load): A text file of emojis into a Python dict</a></span></li><li><span><a href="#2-ways-to-change-types-in-the-values-of-a-dict" data-toc-modified-id="2-ways-to-change-types-in-the-values-of-a-dict-6">2 ways to change types in the values of a <code>dict</code></a></span></li><li><span><a href="#Model-stocks-with-a-custom-class-and-custom-sorting" data-toc-modified-id="Model-stocks-with-a-custom-class-and-custom-sorting-7">Model stocks with a custom class and custom sorting</a></span></li><li><span><a href="#Find-the-largest-sum-with-a-constraint" data-toc-modified-id="Find-the-largest-sum-with-a-constraint-8">Find the largest sum with a constraint</a></span></li></ul></div>

Lab 4: Unicode, Multiple Ways, and Lambda Functions
----

Learning Outcomes
-----

#### By the end of this notebook, you should be able to:

- Properly handle noisy text data.
- Programmatically create and manipulate dictionaries.
- Use `lambda` functions to a create custom sorting functions.

```python
reset -fs
```

Always a string
------

Write a function that takes in either a Unicode string or bytes string and always returns a Python `str` (Unicode string).

Assume utf-8 encoding for the byte string.

<!-- grade_id: cell-97896abd58369ce6, points: 0, grade: false, solution: true -->
```python
from typing import Union

def to_str(s: Union[str, bytes]) -> str:
    
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-26c534b4d0716712, points: 2, grade: true, solution: false -->
```python
"""
Test code to_str function.
2 points.
This cell should NOT give any errors when it is run.
"""

assert type(to_str("hi")) == str
assert to_str("hi") == "hi"

assert type(to_str(b"\xe2\x82\xac20")) == str
assert to_str(b"\xe2\x82\xac20") == "â‚¬20"
```

Messed up content
----

I created a website with Japanese text. However, when loaded into Python it is not correct.

Write a function that changes the content to be correct.

```python
url = 'http://freetexthost.com/01f6ypyoaz'
```

```python
from IPython.display import IFrame

IFrame(url,
      width=800,
      height=600)
```

Output:
```
<IPython.lib.display.IFrame at 0x1095a20b8>
```

```python
# Make sure requests package is installed
try:
    import requests
except ImportError:
    import pip
    import sys
    import subprocess
    subprocess.check_call([sys.executable, '-m', 'pip', 'install', 'requests'])
    import requests
```

```python
# Get the data from the website
content = requests.get(url).content
# Visually inspect it
content[2465:2600]
```

Output:
```
b'\n\n\xe3\x81\xab\xe3\x81\xaa \xe7\x8f\xa6\xe9\xa5\xa6\xe3\x83\x95\xe3\x82\xa9 \xe4\xa9\xa8\xe3\x81\xbd\xe7\x8f\xa6\xe9\xa5\xa6\xe3\x83\x95\xe3\x82\xa9 \xe7\xaa\xaf\xe3\x82\x8a\xe3\x82\x85\xe3\x82\xa6\xe7\xa9\x9f\t\xe3\x81\xa6 \xe3\x82\x8f\xe3\x82\xac\xe8\xa7\x9c\xe3\x83\x95\xe3\x82\xa7\xe5\xab\xa7, \xe3\x81\x92\xe6\x8b\xa5 \xe4\xa5\x9c\xe9\xa7\xba\xe3\x82\xb3\xe3\x82\xb0\xe3\x82\xa1\xe8\xb0\xaf \xe9\xa4\x9a\xe5\xb6\xa3\xe5\x8d\xa4\xe3\x82\x9e\xe8\x9f\xa4 \xe3\x82\x90\xe3\x82\xa4\xe3\x82\x92, \xe3'
```

<!-- grade_id: cell-70cc802ecf691ddf, points: 0, grade: false, solution: true -->
```python
def fix(content):
    
    # YOUR CODE HERE
    raise NotImplementedError()

    return content
```

<!-- grade_id: cell-69fa9e350dbc9c0c, points: 2, grade: true, solution: false -->
```python
"""
Test code for fix function.
2 points.
This cell should NOT give any errors when it is run.
"""

fixed_content = fix(content)
assert fixed_content[2467:2468] == 'ã«'
assert fixed_content[2465:2600] == '\n\nã«ãª ç¦é¥¦ãƒ•ã‚© ä©¨ã½ç¦é¥¦ãƒ•ã‚© çª¯ã‚Šã‚…ã‚¦ç©Ÿ\tã¦ ã‚ã‚¬è§œãƒ•ã‚§å«§, ã’æ‹¥ ä¥œé§ºã‚³ã‚°ã‚¡è°¯ é¤šå¶£å¤ã‚žèŸ¤ ã‚ã‚¤ã‚’, ãƒ•ã‚©æ°© ã‚¤ã‚’ç„¨ä¤¤ãƒŸ ãã‚‡ãã‚…è¦µã‚‰ä¤Ž ã«ã‚ƒç›¨ç©ƒã¢ã‚ƒã¶ é¨Œäª¦ä»£ ç ¥ãƒ„ã‚© ã‚¤ã‚’ç„¨ä¤¤ãƒŸ ä‹¨è¥Šåƒ¯ã«ã‚…ã˜ã‚‡ é¤šå¶£å¤ã‚žèŸ¤ è‹¯æ¥¥ã‚Š ç‹¨ã‚­ãƒ¥å£ ã³æ„¦ã«ãªç§š å°¦è§£ã›¥ç¥§é¨ æ¨šé§Œå¨¦ç¦žç°¯ '
```

ELT (Extract, Transform, Load): A text file of emojis into a Python dict
----

1. Read (extract) `emoji_table.txt`.
1. Munge (transform) the data to clean it up.
1. Put (load) into a Python `dict`.

You are not required to write any functions. You can just code in the cell.

```python
# Make sure ftfy package is installed
try:
    from ftfy import explain_unicode # ðŸ‘ˆ Hint
except ImportError:
    import pip
    import sys
    import subprocess
    subprocess.check_call([sys.executable, '-m', 'pip', 'install', 'ftfy'])
    from ftfy import explain_unicode # ðŸ‘ˆ Hint
```

<!-- grade_id: cell-dcc970f90a0763fb, points: 0, grade: false, solution: true -->
```python
emojis = {}

# YOUR CODE HERE
raise NotImplementedError()

```

<!-- grade_id: cell-6c47ef7e929e6bba, points: 5, grade: true, solution: false -->
```python
"""
Test code for ETL emojis.
5 points.
This cell should NOT give any errors when it is run.
"""
assert emojis['a_smiley1'] == 'ðŸ˜„'
assert emojis['a_smiley2'] == 'ðŸ˜ƒ'
assert emojis['e_family2'] == 'ðŸ‘¨ðŸ‘©ðŸ‘§ðŸ‘¦'
assert emojis['d_ok'] == 'ðŸ‘Œ'
assert emojis['d_poop'] == 'ðŸ’©'
```

2 ways to change types in the values of a `dict`
------

During data a processing, you noticed that the values of a dictionary are being stored as `str`. They should be `int`. 

Convert all type of each item in values to `int`. Do it two ways:

1. Procedurally
1. Functionally

```python
original =  {'A': ['1', '2', '3', '4'],
             'B': ['5', '6', '7', '8'],
             'C': ['9', '10', '11', '12'],}
```

<!-- grade_id: cell-af56ecdc4fc62f06, points: 0, grade: false, solution: true -->
```python
# Procedurally change the type of all values.

updated = {}
for key, value in original.items():
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
```

<!-- grade_id: cell-778046ec5d9da900, points: 2, grade: true, solution: false -->
```python
"""
Test code for procedurally changing the type of all values.
2 points.
This cell should NOT give any errors when it is run.
"""

assert updated == {'A': [1, 2, 3, 4], 'B': [5, 6, 7, 8], 'C': [9, 10, 11, 12]}
```

```python
# Reset updated dict
updated = {}
```

<!-- grade_id: cell-2d7fbac124b90859, points: 0, grade: false, solution: true -->
```python
# Functionally change the type of all values (can be done with a one-liner).

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-79ca20d7c7c96ad5, points: 4, grade: true, solution: false -->
```python
"""
Test code for functionally changing the type of all values.
4 points.
This cell should NOT give any errors when it is run.
"""

assert updated == {'A': [1, 2, 3, 4], 'B': [5, 6, 7, 8], 'C': [9, 10, 11, 12]}
```

Model stocks with a custom class and custom sorting
-----

You are managing a small portfolio of stocks and want to use Python to help you.

1. Create a dataclass to store the relevant features of your stocks.
1. Sort portfolio by price from smallest-to-largest using a `lambda` function.

<!-- grade_id: cell-afaa9e8c7bf6c512, points: 0, grade: false, solution: true -->
```python
from dataclasses import dataclass

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-1a10058cc090323a, points: 1, grade: true, solution: false -->
```python
"""
Test code for Stock dataclass.
1 point.
This cell should NOT give any errors when it is run.
"""

uber_stock  = Stock(name='UBER', shares=19, price=44.52)
apple_stock = Stock(name='AAPL', shares=100, price=207.74)
berk_stock  = Stock(name='BRK-A', shares=1, price=315_000.00)
fb_stock    = Stock(name='FB', shares=435, price=199.75)

portfolio = [uber_stock, apple_stock, berk_stock, fb_stock]
assert portfolio == [Stock(name='UBER', shares=19, price=44.52),
                     Stock(name='AAPL', shares=100, price=207.74),
                     Stock(name='BRK-A', shares=1, price=315000.0),
                     Stock(name='FB', shares=435, price=199.75)]
```

<!-- grade_id: cell-14ec9cbcb47be819, points: 0, grade: false, solution: true -->
```python
# Sort profolio by price from largest-to-smallest. Use list.sort with a lambda function

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f076863f6e489107, points: 3, grade: true, solution: false -->
```python
"""
Test code for sorting from largest-to-smallest.
3 point.
This cell should NOT give any errors when it is run.
"""

assert portfolio == [Stock(name='BRK-A', shares=1, price=315000.0),
                     Stock(name='AAPL', shares=100, price=207.74),
                     Stock(name='FB', shares=435, price=199.75),
                     Stock(name='UBER', shares=19, price=44.52)]
```

Find the largest sum with a constraint
-----

Given a list of integers, find the largest sum with the constraint that you can not add two adjacent numbers.

Your code should be efficient. Minimize the number of passes through the data. If your function times-out during autograding, you'll lose the points. 

<!-- grade_id: cell-71626a6b8bdf720b, points: 0, grade: false, solution: true -->
```python
from typing import List

def funky_sum(nums: List[int]) -> int:
    
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-0f8b5ca7275ae799, points: 3, grade: true, solution: false -->
```python
"""
Test code for funky_sum function.
3 points for correctness.
This cell should NOT give any errors when it is run.
"""

                                          # Indicator vector
assert funky_sum([1, 2, 3, 1])     ==  4  # [1, 0, 1, 0]
assert funky_sum([2, 7, 9, 3, 1])  == 12  # [1, 0, 1, 0, 1]
assert funky_sum([2, 1, 1, 2])     ==  4  # [1, 0, 0, 1]
```

<!-- grade_id: cell-edfa95f882c99fee, points: 3, grade: true, solution: false -->
```python
"""
Test code for funky_sum function.
3 points for scaling.
This cell should NOT give any errors when it is run.
"""

n = 100_000
assert funky_sum([2, 7, 9, 3, 1]*n) == 1100001
assert funky_sum([2, 1, 1, 2]*n) == 300001
```

<br>
<br> 
<br>

----