<h1>Table of Contents<span class="tocSkip"></span></h1>
<div class="toc"><ul class="toc-item"><li><span><a href="#Lab-1:-Fun-With-Functions-" data-toc-modified-id="Lab-1:-Fun-With-Functions--1">Lab 1: Fun With Functions </a></span></li><li><span><a href="#Learning-Outcomes" data-toc-modified-id="Learning-Outcomes-2">Learning Outcomes</a></span><ul class="toc-item"><li><ul class="toc-item"><li><span><a href="#By-the-end-of-this-notebook,-you-should-be-able-to:" data-toc-modified-id="By-the-end-of-this-notebook,-you-should-be-able-to:-2.0.1">By the end of this notebook, you should be able to:</a></span></li></ul></li></ul></li><li><span><a href="#Find-the-Runner-Up-Score!" data-toc-modified-id="Find-the-Runner-Up-Score!-3">Find the Runner-Up Score!</a></span></li><li><span><a href="#Binomial-coefficient-" data-toc-modified-id="Binomial-coefficient--4">Binomial coefficient </a></span></li><li><span><a href="#Funny-sort" data-toc-modified-id="Funny-sort-5">Funny sort</a></span></li><li><span><a href="#Lower-Triangular-Matrix" data-toc-modified-id="Lower-Triangular-Matrix-6">Lower Triangular Matrix</a></span></li><li><span><a href="#Can-I-sleep-in?" data-toc-modified-id="Can-I-sleep-in?-7">Can I sleep-in?</a></span></li><li><span><a href="#Happy-Number" data-toc-modified-id="Happy-Number-8">Happy Number</a></span></li><li><span><a href="#Greatest-Common-Divisor" data-toc-modified-id="Greatest-Common-Divisor-9">Greatest Common Divisor</a></span></li><li><span><a href="#Find-the-largest-sum-for-a-contiguous-sublist" data-toc-modified-id="Find-the-largest-sum-for-a-contiguous-sublist-10">Find the largest sum for a contiguous sublist</a></span></li></ul></div>

Lab 1: Fun With Functions 
-----

Learning Outcomes
-----

#### By the end of this notebook, you should be able to:

- Write Python functions to solve data science related problems.
- Submit lab for autograding.

Do not import anything additionally. Use only built-in Python.

```python
reset -fs
```

Find the Runner-Up Score!
-----

Tackle [Find the Runner-Up Score!](https://www.hackerrank.com/contests/msds-incoming-student/challenges/find-second-maximum-number-in-a-list) from the HackerRank challenges.

Try to do it in one-liner with built-in functions.

Hints:

- Use two different built-ins functions as a functional sandwich.
- Using keyword arguments will make your code easier to read and reason about.

<!-- grade_id: cell-ee8bfea85d58df23, points: 0, grade: false, solution: true -->
```python
from typing import List

def find_runner_up(nums: List) -> int:
    
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: cell-5eaaa9de6e729f4a, points: 1, grade: true, solution: false -->
```python
"""
Test code for find_runner_up function.
1 points.
This cell should NOT give any errors when it is run.
"""

assert find_runner_up([1, 2, 3]) == 2
assert find_runner_up([2, 3, 6, 6, 5]) == 5
assert find_runner_up([1, 3, 3, 3, 2]) == 2
```

Binomial coefficient 
-----

Write a function that returns the number of ways to choose k items from n items without repetition and without order.

[Binomial coefficient](https://en.wikipedia.org/wiki/Binomial_coefficient) is mathematically equal to the expression $ \frac{n!}{(k! (n - k)!)}$

<!-- grade_id: cell-d89c48480247256d, points: 0, grade: false, solution: true -->
```python
from math import factorial # Factorial is the product of an integer and all the integers below

def n_combinations(n:int, k:int) -> int: 
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return result
```

<!-- grade_id: cell-be8a34383a4857ce, points: 1, grade: true, solution: false -->
```python
"""
Test code for n_combinations function.
1 points.
This cell should NOT give any errors when it is run.
"""

from scipy.special import binom

assert n_combinations(n=8, k=0)  == binom(8, 0)  == 1
assert n_combinations(n=8, k=1)  == binom(8, 1)  == 8
assert n_combinations(n=8, k=2)  == binom(8, 2)  == 28
assert n_combinations(n=8, k=3)  == binom(8, 3)  == 56
assert n_combinations(n=8, k=4)  == binom(8, 4)  == 70
assert n_combinations(n=8, k=5)  == binom(8, 5)  == 56
assert n_combinations(n=8, k=6)  == binom(8, 6)  == 28
assert n_combinations(n=8, k=7)  == binom(8, 7)  == 8
assert n_combinations(n=8, k=8)  == binom(8, 8)  == 1
assert n_combinations(n=39, k=5) == binom(39, 5) == 575_757
```

Funny sort
-----

Sort a list of numbers so evens are one side and odds are on the other.

Implement it from scratch. That means looping through the data and writing the comparison logic. Do not use the `sorted` or `list.sort`.

This is a simplified example of the data science problem of sorting people by acquisition channel. At a previous company, my data team encoded a lot information into user id, including acquisition channel.

<!-- grade_id: cell-a1f2c31dbab1ae70, points: 0, grade: false, solution: true -->
```python
from typing import List

def sort_evens_odds(numbers: List[int]):
    
    # YOUR CODE HERE
    raise NotImplementedError()
    
    return numbers
```

<!-- grade_id: cell-da870659a34de14f, points: 2, grade: true, solution: false -->
```python
"""
Test code for sort_evens_odds function.
2 points.
This cell should NOT give any errors when it is run.
"""

numbers = [1, 2, 2, 3, 4, 4, 5, 6]
assert sort_evens_odds(numbers) == [2, 2, 4, 4, 6, 1, 3, 5]
```

Lower Triangular Matrix
------

Add code to the inner loop so it sets the upper left triangle of the matrix to zero, aka lower triangular matrix form.

[Lower triangular matrix](https://en.wikipedia.org/wiki/Triangular_matrix) is very useful for [LU decomposition](https://en.wikipedia.org/wiki/LU_decomposition).

<!-- grade_id: cell-ed82dda441f186f3, points: 0, grade: false, solution: true -->
```python
from typing import List

def zero_out_triangle(matrix: List[List[int]]) -> List[List[int]]:
    
    for i, row in enumerate(matrix):
        for j, col in enumerate(row):
            
            # YOUR CODE HERE
            raise NotImplementedError()
            
    return matrix

```

```python
# Visually inspect your answer
print(*zero_out_triangle( [[1, 2, 3], [4, 5, 6], [7, 8, 9]]), sep='\n')
```

<!-- grade_id: cell-5f7404752ade9fa3, points: 2, grade: true, solution: false -->
```python
"""
Test code for zero_out_triangle function.
2 points.
This cell should NOT give any errors when it is run.
"""

assert zero_out_triangle( [[1, 2], [3, 4]]) ==  [[1, 0],
                                                 [3, 4]]
assert zero_out_triangle( [[1, 2, 3], [4, 5, 6], [7, 8, 9]]) == [[1, 0, 0], 
                                                                 [4, 5, 0], 
                                                                 [7, 8, 9]]
```

Can I sleep-in?
------

Write a function that tells you if you can sleep-in.

You can sleep-in if it is not a weekday or on vacation.

The function will take two parameters:

1. `weekday` is True if it is a weekday
1. `vacation` is True if you are on vacation

The function will return `True` if you can sleep-in or `False` if you can not sleep-in.

This is an example of a truth-table logic for whether a user would get a feature that is currently being A/B tested.

<!-- grade_id: cell-245d0953a598a6fe, points: 0, grade: false, solution: true -->
```python
def sleep_in(weekday: bool, vacation: bool) -> bool:

    # YOUR CODE HERE
    raise NotImplementedError()
    
```

<!-- grade_id: cell-771475a0ff7684b3, points: 2, grade: true, solution: false -->
```python
"""
Test code for sleep_in function.
2 points.
This cell should NOT give any errors when it is run.
"""

assert sleep_in(weekday=False, vacation=False)       == True
assert sleep_in(weekday=True,  vacation=False)       == False
assert sleep_in(weekday=False, vacation=True)        == True
assert sleep_in(weekday=True,  vacation=True)        == True
```

Happy Number
------

This is not a Data Science problem. However, it is interesting numerology!

A happy number is a number defined by the following process: Starting with any positive integer, replace the number by the sum of the squares of its digits, and repeat the process until the number equals 1 (where it will stay), or it loops endlessly in a cycle which does not include 1. 

Those numbers for which this process ends in 1 are happy numbers.

19 is a happy number:

$1^2 + 9^2 = 82$

$8^2 + 2^2 = 68$

$6^2 + 8^2 = 100$

$1^2 + 0^2 + 0^2 = 1$

An unhappy number is a number that is not happy. In other words, all unhappy numbers eventually start repeating a numbers in a cycle:

$4 → 16 → 37 → 58 → 89 → 145 → 42 → 20 → 4$

You should notice when a number repeats and stop checking. Then, declare the initial number as unhappy.

Your solution should be efficient. If your function it times out, you won't get the points.

<!-- grade_id: cell-e51c96c5ab7d983e, points: 0, grade: false, solution: true -->
```python
def is_happy(n: int) -> bool:
    
    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: cell-7c51783ddf1cd534, points: 3, grade: true, solution: false -->
```python
"""
Test code for is_happy function.
3 points.
This cell should NOT give any errors when it is run.
"""

assert is_happy(1)
assert is_happy(7)
assert is_happy(10)
assert is_happy(19)

assert not is_happy(2)
assert not is_happy(3)
assert not is_happy(8)
```

Greatest Common Divisor
-----

The greatest common divisor (GCD) is the largest natural number $d$ that divides both $a$ number and $b$ number without a remainder. For example, the GCD the fraction 8 and 12 is 4: 

$$8/4 = 2$$ 
$$12/4 = 3$$

Write the naive, brute-force implementation - Iterate through all integers from 1 to `max(a, b)` to find the largest common divisor of the given numbers.

<!-- grade_id: cell-a98050a4af0aef8b, points: 0, grade: false, solution: true -->
```python
def gcd(a: int, b: int) -> int:

    gcd = 1

    # YOUR CODE HERE
    raise NotImplementedError()
    
    return gcd
```

<!-- grade_id: cell-ea9ea7a18818a921, points: 3, grade: true, solution: false -->
```python
"""
Test code for gcd function.
3 points.
This cell should NOT give any errors when it is run.
"""

assert gcd(1, 1) == 1
assert gcd(1, 1) == 1
assert gcd(3, 9) == 3
assert gcd(12, 8) == 4
assert gcd(12, 24) == 12
assert gcd(17, 13) == 1
assert gcd(45, -12) == 3
```

Find the largest sum for a contiguous sublist
-----

Given a list of integers, find the largest sum for a contiguous sublist.

For example, you are writing a simple investing bot that can only buy and sell once. You need to find the how good is the best run of a stock. Thus, the function will find the maximum return your simple bot can earn.

Your code should be efficient. Minimize the number of passes through the data. If your function times-out during autograding, you'll lose the points. 

<!-- grade_id: cell-6b624cbb60497282, points: 0, grade: false, solution: true -->
```python
from typing import List

def max_subarrary(nums: List[float]) -> float:

    # YOUR CODE HERE
    raise NotImplementedError()

```

<!-- grade_id: cell-07d9d80d492e82ec, points: 2, grade: true, solution: false -->
```python
"""
Test code for max_subarrary function.
2 points.
This cell should NOT give any errors when it is run.
"""

# Basic test cases
assert max_subarrary(nums=[0, 1, 1]) == 2
assert max_subarrary(nums=[1, -1, 1, 1, -1]) == 2
```

<!-- grade_id: cell-e9edcbf6385057f0, points: 2, grade: true, solution: false -->
```python
"""
Test code for max_subarrary function.
2 points.
This cell should NOT give any errors when it is run.
"""

# Harder test case for logic
assert max_subarrary(nums=[-2, 1, -3, 4, -1, 2, 1, -5, 4]) == 6
```

<!-- grade_id: cell-4eabd17bc0228d19, points: 2, grade: true, solution: false -->
```python
"""
Test code for max_subarrary function.
2 points.
This cell should NOT give any errors when it is run.
"""

# Harder test case for scaling
n_filler = 2_000_000
assert max_subarrary(nums=[-1]*n_filler+[1, 0, 1]+[-1]*n_filler+[1, 0, 1, 1]+[-1]*n_filler) == 3
```

<br>
<br> 
<br>

----