<h1>Table of Contents<span class="tocSkip"></span></h1>
<div class="toc"><ul class="toc-item"><li><span><a href="#Lab-3:-Recursion-and-OOP" data-toc-modified-id="Lab-3:-Recursion-and-OOP-1">Lab 3: Recursion and OOP</a></span></li><li><span><a href="#Learning-Outcomes" data-toc-modified-id="Learning-Outcomes-2">Learning Outcomes</a></span><ul class="toc-item"><li><ul class="toc-item"><li><span><a href="#By-the-end-of-this-notebook,-you-should-be-able-to:" data-toc-modified-id="By-the-end-of-this-notebook,-you-should-be-able-to:-2.0.1">By the end of this notebook, you should be able to:</a></span></li></ul></li></ul></li><li><span><a href="#Exponentiation-with-recursion" data-toc-modified-id="Exponentiation-with-recursion-3">Exponentiation with recursion</a></span></li><li><span><a href="#Create-a-Fibonacci-sequence-with-recursion" data-toc-modified-id="Create-a-Fibonacci-sequence-with-recursion-4">Create a Fibonacci sequence with recursion</a></span></li><li><span><a href="#Greatest-common-divisor-(GCD)-with-recursion" data-toc-modified-id="Greatest-common-divisor-(GCD)-with-recursion-5">Greatest common divisor (GCD) with recursion</a></span></li><li><span><a href="#Len(gth)-of-a-list-with-recursion" data-toc-modified-id="Len(gth)-of-a-list-with-recursion-6">Len(gth) of a list with recursion</a></span></li><li><span><a href="#Reverse-a-string-with-recursion" data-toc-modified-id="Reverse-a-string-with-recursion-7">Reverse a string with recursion</a></span></li><li><span><a href="#Make-a-data-class-for-server-log-data" data-toc-modified-id="Make-a-data-class-for-server-log-data-8">Make a data class for server log data</a></span></li></ul></div>

Lab 3: Recursion and OOP
-----

Learning Outcomes
-----

#### By the end of this notebook, you should be able to:

- Solve a variety of problems with recursion.
- Create a data class to model a specific type of data.

```python
reset -fs
```

Exponentiation with recursion
---

Write an [exponentiation](https://en.wikipedia.org/wiki/Exponentiation) function that is recursive.

<!-- grade_id: cell-a0410474f2612710, points: 0, grade: false, solution: true -->
```python
def exp(b: int, x: int) -> int:
    
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-c552409509d50242, points: 1, grade: true, solution: false -->
```python
"""
Test code for exp function.
1 points for correctness.
This cell should NOT give any errors when it is run.

Note: Your `exp` is very in the range of valid inputs. 
      Do not be concerned about negative or fractional exponents.
"""

assert exp(2, 0) == 1
assert exp(2, 1) == 2
assert exp(2, 2) == 4
assert exp(2, 3) == 8
assert exp(2, 4) == 16
assert exp(2, 5) == 32
assert exp(3, 0) == 1
assert exp(3, 1) == 3
assert exp(3, 2) == 9
assert exp(3, 3) == 27
```

<!-- grade_id: cell-c79dd3bbf9a254d1, points: 3, grade: true, solution: false -->
```python
"""
Test code for exp function.
3 points for using recursion.
This cell should NOT give any errors when it is run.
"""

from bdb import Bdb
import sys

class RecursionDetected(Exception):
    pass

class RecursionDetector(Bdb):
    "From: https://stackoverflow.com/questions/36662181/is-there-a-way-to-check-if-function-is-recursive-in-python"
    def do_clear(self, arg):
        pass

    def __init__(self, *args):
        Bdb.__init__(self, *args)
        self.stack = set()

    def user_call(self, frame, argument_list):
        code = frame.f_code
        if code in self.stack:
            raise RecursionDetected
        self.stack.add(code)

    def user_return(self, frame, return_value):
        self.stack.remove(frame.f_code)

def test_recursion(func):
    detector = RecursionDetector()
    detector.set_trace()
    try:
        func()
    except RecursionDetected:
        return True
    else:
        return False
    finally:
        sys.settrace(None)
        
# Test if function is recursive
assert test_recursion(lambda: exp(2, 5))
```

Create a Fibonacci sequence with recursion
----

Parts:

1. Load Fibonacci sequence  for tests
1. Write Fibonacci sequence that is recursive

__Part 1: Load fib numbers__

Load `fib_seq.txt` into a variable named `fibs_actual`.

Do __not__ write a function. Just write the code in the cell below.

<!-- grade_id: cell-7664463a45dc6415, points: 0, grade: false, solution: true -->
```python

# YOUR CODE HERE
raise NotImplementedError()

```

<!-- grade_id: cell-99764163a03433cf, points: 0, grade: true, solution: false -->
```python
"""
Test code for loading fib_seq.txt.
2 points
This cell should NOT give any errors when it is run.
"""

assert len(fibs_actual) == 30
assert fibs_actual[:3] == [0, 1, 1]
assert fibs_actual[-3:] == [196418, 317811, 514229]
```

__Part 2: Generate Fibonacci sequence with recursion__


<!-- grade_id: cell-2cea8f1aacb0b678, points: 0, grade: false, solution: true -->
```python
def fib(n: int) -> int:

    # YOUR CODE HERE
    raise NotImplementedError()
    
```

<!-- grade_id: cell-29b00957b6f83a1d, points: 1, grade: true, solution: false -->
```python
"""
Test code for fib function.
1 points for correctness.
This cell should NOT give any errors when it is run.
"""

fib_generated = [fib(n) for n in range(30)] 
assert fib_generated == fibs_actual
```

<!-- grade_id: cell-23ef20ce21b2010f, points: 3, grade: true, solution: false -->
```python
"""
Test code for fib function.
3 points for using recursion.
This cell should NOT give any errors when it is run.
"""

assert test_recursion(lambda: fib(10))
```

Greatest common divisor (GCD) with recursion
----

Write a function to find the the greatest common divisor for two numbers with recursion.

This is also called Euclidean Algorithm. The Greek mathematician Euclid described this algorithm approx. 300 BCE and still being used in the field of number theory and cryptography.

<!-- grade_id: cell-b3a108fc09a64bcd, points: 0, grade: false, solution: true -->
```python
def gcd(a: int, b: int) -> int:
    
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-cd5d0a53c67f6a40, points: 1, grade: true, solution: false -->
```python
"""
Test code for gcd function.
1 point for correctness.
This cell should NOT give any errors when it is run.
"""

assert gcd(1, 1) == 1
assert gcd(1, 1) == 1
assert gcd(3, 9) == 3
assert gcd(12, 8) == 4
assert gcd(12, 24) == 12
assert gcd(17, 13) == 1
assert gcd(45, -12) == 3
```

<!-- grade_id: cell-37055e2726fc945f, points: 3, grade: true, solution: false -->
```python
"""
Test code for gcd function.
3 points for using recursion.
This cell should NOT give any errors when it is run.
"""

assert test_recursion(lambda: gcd(45, -12))
```

Len(gth) of a list with recursion
----

Find a recursive function that returns the number of items in a list.

<!-- grade_id: cell-e50f07139b923d61, points: 0, grade: false, solution: true -->
```python
def len_list(a: list) -> int:
    
    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-ec281ff9e30dab21, points: 1, grade: true, solution: false -->
```python
"""
Test code for len_list function.
1 point for correctness.
This cell should NOT give any errors when it is run.
"""

assert len_list([]) == len([]) == 0
assert len_list([None]) == len([None]) == 1
assert len_list([1, 2, 3]) == len([1, 2, 3]) == 3
```

<!-- grade_id: cell-79b6255cc88b4741, points: 3, grade: true, solution: false -->
```python
"""
Test code for len_list function.
3 points for using recursion.
This cell should NOT give any errors when it is run.
"""

assert test_recursion(lambda: len_list([1, 2, 3]))
```

Reverse a string with recursion
----

Brian once had a technical interview which was solely "Brainstorm and implement as many ways as possible to reverse a string."

He got tripped up trying to implement reversing a string with recursion. 

Be better than Brian.

<!-- grade_id: cell-84f0c22e48d77c38, points: 0, grade: false, solution: true -->
```python
def reverse_string(string: str) -> str:

    # YOUR CODE HERE
    raise NotImplementedError()
```

<!-- grade_id: cell-e23e8f3c94043460, points: 1, grade: true, solution: false -->
```python
"""
Test code for reverse_string function.
1 point for correctness.
This cell should NOT give any errors when it is run.
"""

string = "redrum"
assert reverse_string(string) == "murder" == "".join(reversed(string)) == string[::-1]
```

<!-- grade_id: cell-000f3bf88e1d4533, points: 3, grade: true, solution: false -->
```python
"""
Test code for reverse_string function.
3 points for using recursion.
This cell should NOT give any errors when it is run.
"""

assert test_recursion(lambda: reverse_string("redrum"))
```

Make a data class for server log data
------

Data class allow us to store data in class. A class is more human readable each field has a name and easier to access because the `.` interface.

Server log files are a raw, unfiltered look at the traffic to a website. They are the text files stored on a webserver. Every time anything (mostly browsers) request anything (mostly urls) from your server, the server adds a line in the log. The logs also store the errors and the attacks on a webserver.

Define a data class called `ApacheLog` with following attributes:

- ip
- datestamp
- status_code

Since ApacheLog is a data class, ApacheLog will have no methods. 

Examples of real Apache web server logs are [here](http://www.monitorware.com/en/logsamples/apache.php).

For this assignment, just use generic built-in types for type hints. In real code, we could be more specific with type hints. 

<!-- grade_id: cell-d20f2c6bd85416c4, points: 0, grade: false, solution: true -->
```python
from dataclasses import dataclass

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-f967452b7b148e5e, points: 1, grade: true, solution: false -->
```python
"""
Test code for ApacheLog class.
1 point
This cell should NOT give any errors when it is run.
"""

l = ApacheLog(ip='192.168.0.1',
              datestamp='07/Mar/2019:16:05:49 -0800',
              status_code=200,
              )
attributes =[a for a in dir(l) if not a.startswith("__")]

assert isinstance(l, ApacheLog)
assert set(attributes) == {'datestamp', 'ip', 'status_code'}
assert '__dataclass_fields__' in dir(l)
```

<br>
<br> 
<br>

----